Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = "NONE"
```

# Python Programming (very fast) Introduction

## Importing modules
Python has an extensive standard library (https://docs.python.org/3/library/) to perform several task in optimized and easy-to-use way. For basic mathematical operations you can just use the normal operators, but for other standard functions, like `sin`, you could import the standard `math` module, or , much better, use either `numpy` or `scipy` modules
```python
import math as m # imports the math module with an alias m
# import math # imports math module but you need to use math.sin or math.cos
# from math import * # DO NOT do this: exposes all inners from math

x = 5.2
y = m.sin(x)
print(x, y)
```

<div class="alert alert-block alert-info">
<b>Exercise:</b> Copy and execute the previous code in the following cell (use shift+enter to execute) 
</div>

<!-- grade_id: cell-1d206d3930590507, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<div class="alert alert-block alert-info">
<b>Exercise:</b> Now replace math by numpy and execute it. Is there any difference?
</div>

<!-- grade_id: cell-d2ac73e7587d20db, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Looking for help
What are the units of the argument for the `sin` function? Are there any other functions provided by the `numpy` module? Check the documentation using either google or, inside a notebook
```
help(np.sin)
help(np)
np.sin?
```

```python
help(np)
```

## Printing and reading to and from the screen

You have already seen how to print to the screen. To read, you can use the function call ```input```. Try it:

```python
var = input("Please write something :\t")
print("You have written : " + var)
print(f"You have written : {var}") # uses a f-string: https://realpython.com/python-f-strings/
```

You can also format the printing string, like you have done in c/c++ . For example, you can write something like

```python
x = 5
y = 2.34
print("x = %04d \t y = %.1f \n y=%20.16e\n"%(x, y, y))
print(f"x = {x:04d} \t y = {y:.1f} \n y={y:20.16e}")
```

Note: if you need to manipulate strings and numbers, you should cast the number to string by using ```str```, as in ```str(x)``` .

<div class="alert alert-block alert-info">
<b>Exercise:</b> Print the sum of two numbers with 5 decimals, filled with 0, and in field of 10 spaces 
</div>

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
x = 32.6
y = -4.7
# YOUR CODE HERE
raise NotImplementedError()
```

## Selection structures : `if`, `if else`

This is a simple example on how to use if to select from several options

```python
m = 5
if m < 1:
    print (f"{m} is smaller than 1")
elif 1 <= m <= 4:
    print (f"{m} is in between 1 and 4")
else :
    print (f"{m} is larger than 4")
```

Output:
```
5 is larger than 4
```

You can also compare values directly, or even compare several values

```python
print(2 == 3)
print(2 == 2 and 3 > 9)
print([2, 3] == [3, 2])
print([2, 3] == [2, 2])
print([2, 3] == [2, 3])
```

Output:
```
False
False
False
False
True
```

<div class="alert alert-block alert-info">
<b>Exercise:</b> Read a number and indicate if the number is even or odd 
</div>

<!-- grade_id: cell-8ccf3fbf15c165e4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()
m = int(input("Por favor ingrese un numero:"))
if m%2 == 0 :
    print(f"{m} es par")
else:
    print(f"{m} es impar")
```

Output:
```
Por favor ingrese un numero: 1232133
1232133 es impar
```

## Loop structures: for and while

```python
N = 5
i = 0
while i < N:
    print (i, end=" ")
    i = i+1
print("End")
print(f"{i=}")
```

Output:
```
0,1,2,3,4,End
i=5
```

```python
for ii in range(1, 10):
    print(f"{ii}", end=" ")
```

Output:
```
1 2 3 4 5 6 7 8 9
```

```python
name = "My name" # a string is an array for characters
for c in name:
    print (c, end="-")
print()
print(name[0])
print(name[-1])
```

Output:
```
M-y- -n-a-m-e-
M
e
```

<div class="alert alert-block alert-info">
<b>Exercise:</b> Print values from 0.0, to 1.0 (exclusive), in steps of 0.05 
</div>

<!-- grade_id: cell-a83db8b3e865fef8, points: 0, grade: false, solution: true -->
```python
x = 0.0
while x < 1.0:
    print(x)
    x = x + 0.05
# YOUR CODE HERE
#raise NotImplementedError()
```

Output:
```
0.0
0.05
0.1
0.15000000000000002
0.2
0.25
0.3
0.35
0.39999999999999997
0.44999999999999996
0.49999999999999994
0.5499999999999999
0.6
0.65
0.7000000000000001
0.7500000000000001
0.8000000000000002
0.8500000000000002
0.9000000000000002
0.9500000000000003
```

<div class="alert alert-block alert-info">
<b>Exercise:</b> Now use numpy to do the same (check for the linspace function) 
</div>

<!-- grade_id: cell-4d62b200c665b8ab, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()
import numpy as np
xmin = 0.0
xmax = 1.0
deltax = 0.05
N = int((xmax-xmin)/deltax)
print(np.linspace(xmin, xmax, N, endpoint=False))
print(np.arange(0.0, 1.0, 0.05))
```

Output:
```
[0.   0.05 0.1  0.15 0.2  0.25 0.3  0.35 0.4  0.45 0.5  0.55 0.6  0.65
 0.7  0.75 0.8  0.85 0.9  0.95]
[0.   0.05 0.1  0.15 0.2  0.25 0.3  0.35 0.4  0.45 0.5  0.55 0.6  0.65
 0.7  0.75 0.8  0.85 0.9  0.95]
```

## How to declare functions
Functions are the next construct necessary for the a good program design. A function should be defined by using the keyword ```def``` .

```python
def my_function():
    print ("I am the function")
    
my_function() # this calls the function    
my_function()
```

Output:
```
I am the function
I am the function
```

```python
%%writefile test.py
# this shows how to pass arguments
def my_function(x:float):
    print (f"Value of argument is {x}")
    y = 2.0*x
    
#my_function() # error, requires argument
my_function(-3)
my_function(0.9)
my_function("Hola")
```

Output:
```
Writing test.py
```

```python
# this shows how to pass arguments, and assign default values
def my_function(x = 0):
    """
    kljshadasdhaasd sa
    asd asd sa asd
    x : asdsa
    output: 
    """
    print (f"Value of argument is {x}")

my_function() # no longer an error
my_function(-3)
```

Output:
```
Value of argument is 0
Value of argument is -3
```

```python
# this shows how to pass arguments, assign default values, and also calls by named arguments
def my_function(x = 0, y = 1):
    print ("Calling my_function")
    print ("x " + str(x))
    print ("y " + str(y))
    
my_function() # no longer an error
my_function(-3, 2) # x = -3, y = -2
my_function(y=-3, x=2) # x = 2, y = -3
```

# Exercises

## Convert from Farenheit to Celcius
Write a function that receives a temperature in F and prints it in C 

<!-- grade_id: cell-6d72db8e857ab57e, points: 0, grade: false, solution: true -->
```python
def FtoC(f):
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
Cell In [44], line 3
    raise NotImplementedError()
    ^
IndentationError: expected an indented block after function definition on line 1
```

<!-- grade_id: cell-0764b89cea883781, points: 0, grade: true, solution: false -->
```python
assert FtoC(0) == "The temperature 0 F is equivalent to -17.78 C"
assert FtoC(10.3) == "The temperature 10.3 F is equivalent to -12.06 C"
### BEGIN HIDDEN TEST
assert FtoC(-10.3) == "The temperature -10.3 F is equivalent to -23.50 C"
### END HIDDEN TEST
```

Output:
```
Unexpected exception formatting exception. Falling back to standard exception
Traceback (most recent call last):
  File "/Users/oquendo/opt/miniforge3/envs/herrcomp/lib/python3.10/site-packages/IPython/core/interactiveshell.py", line 3378, in run_code
    exec(code_obj, self.user_global_ns, self.user_ns)
  File "/var/folders/rw/w_6hzhcn4wsdpw6bk8bzr14w0000gn/T/ipykernel_32903/2503453734.py", line 1, in <module>
    assert FtoC(0) == "The temperature 0 F is equivalent to -17.78 C"
NameError: name 'FtoC' is not defined

During handling of the above exception, another exception occurred:

Traceback (most recent call last):
  File "/Users/oquendo/opt/miniforge3/envs/herrcomp/lib/python3.10/site-packages/IPython/core/interactiveshell.py", line 1997, in showtraceback
    stb = self.InteractiveTB.structured_traceback(
  File "/Users/oquendo/opt/miniforge3/envs/herrcomp/lib/python3.10/site-packages/IPython/core/ultratb.py", line 1112, in structured_traceback
    return FormattedTB.structured_t
```

```python
# Listas
data = [1, 2, 3, "asdsa"]
print(data[1])
# Tuplas
tup = (1, 2, 3)
# Dictionarios
mydict = {"valora":454.3454, "valorb":-98.87768}
print(mydict["valora"])
mydict["valora"] = -1.2344343
print(mydict["valora"])
def myfun(param: dict[str, float]):
    print(2*param["valorb"])
myfun(mydict)

```

Output:
```
2
454.3454
-1.2344343
Unexpected exception formatting exception. Falling back to standard exception
Traceback (most recent call last):
  File "/Users/oquendo/opt/miniforge3/envs/herrcomp/lib/python3.10/site-packages/IPython/core/interactiveshell.py", line 3378, in run_code
    exec(code_obj, self.user_global_ns, self.user_ns)
  File "/var/folders/rw/w_6hzhcn4wsdpw6bk8bzr14w0000gn/T/ipykernel_32903/467724190.py", line 13, in <module>
    myfun(data)
  File "/var/folders/rw/w_6hzhcn4wsdpw6bk8bzr14w0000gn/T/ipykernel_32903/467724190.py", line 12, in myfun
    print(2*param["valorb"])
TypeError: list indices must be integers or slices, not str

During handling of the above exception, another exception occurred:

Traceback (most recent call last):
  File "/Users/oquendo/opt/miniforge3/envs/herrcomp/lib/python3.10/site-packages/IPython/core/interactiveshell.py", line 1997, in showtraceback
    stb = self.InteractiveTB.structured_traceback(
  File "/Users/oquendo/opt/miniforge3/envs/herrcomp/
```

## Triangle vertexes
The coordinates for the vertexes of a given triangle are written into a file in the form  \begin{align*}x_1\ \ \  y_1\\x_2\ \ \  y_2\\x_3\ \ \ y_3  \end{align*} Write a program that reads this values, and computes the area of the triangle, defined as $$ A = \frac{1}{2}[x_2y_3 - x_3y_2 - x_1y_3 + x_3y_1 + x_1y_2 - x_2y_1]$$. Use functions.

## RC circuit
Write a Python program that uses a loop to simulate the behavior of an RC circuit. The program should ask the user for the resistance and capacitance values of the circuit and should calculate the voltage across the capacitor as a function of time. The program should then create a plot of the voltage across the capacitor versus time.

## Spring system
Write a Python program that uses a loop to simulate the motion of a mass-spring system. The program should ask the user for the mass of the object, the spring constant, and the initial displacement from equilibrium. The program should then calculate the position of the object as a function of time and create a plot of the position versus time.

## Projectile
Write a Python program that uses a loop to simulate the motion of a projectile under the influence of gravity. The program should ask the user for the initial velocity and launch angle of the projectile, and should calculate the trajectory of the projectile over time. The program should then create a plot of the trajectory of the projectile. Use the Euler integration algorithm to compute the next position and velocity. Compare adding a damping linear force with air. 

## Random numbers
Write a Python program that reads three integers N, lower, upper, and generates a list of N random integers between lower and upper, inclusive. The program should then print out the number of integers in the list that are even and the number of integers that are odd. Finally, the program should create a bar chart that shows the number of even and odd integers in the list. Then create 4 plots with increasing N showing that limit distribution of even and odd numbers.

## Heatmap

Write a Python function that takes an integer n as input and returns a NumPy array of size n x n, where the elements of the array are random integers between lower and upper, which are also arguments of the function. Write a second function that takes the NumPy array as input and creates a heatmap of the array. Use these functions to create a program that generates a mxm NumPy array of random integers and displays it as a heatmap. Use several m


## Fitting with scipy
Here we will learn how to perform a non-linear fit. To do so, we will generate some random data and then fit it. First, create a function that represents the theoretical trend that you want to fit. Now create two numpy arrays, xdata and ydata. For xdata use the linspace function on the range you are interested in. For ydata use the theoretical function and add a gaussian noise for each point. Now use the scipy.curve_fit function to fit the data, and print the fit parameters and their errors. Plot both the data and the fit.

```python

```